## Imports

```python
# pip install --upgrade scikit-learn
```

```python
import time
from IPython.display import clear_output
import numpy    as np
import pandas   as pd
import seaborn  as sb
import matplotlib.pyplot as plt
import sklearn  as skl

from sklearn import pipeline      # Pipeline
from sklearn import preprocessing # OrdinalEncoder, LabelEncoder
from sklearn import impute
from sklearn import compose
from sklearn import model_selection # train_test_split
from sklearn import metrics         # accuracy_score, balanced_accuracy_score, plot_confusion_matrix
from sklearn import set_config

set_config(display='diagram') # Useful for display the pipeline

print("Pandas  ", pd.__version__)
print("Sklearn ", skl.__version__) # Try to use 0.24
```

Output:
```
Pandas   1.2.5
Sklearn  0.24.2
```

## Get the dataset
- **CLOUD = True**: Download dataset from Kaggle. Necesary for cloud enviroments like COLAB. **Specify your [kaggle credentials](https://www.kaggle.com/docs/api)**.
- **CLOUD = False**: Get the dataset from your local machine. **Specify the data path**.

```python
DATA_PATH = "../../Datasets/Tabular/titanic/"
```

## Load data

```python
df      = pd.read_csv(DATA_PATH + "train.csv", index_col='PassengerId')
df_test = pd.read_csv(DATA_PATH + "test.csv",  index_col='PassengerId')

print("Train DataFrame:", df.shape)
print("Test DataFrame: ", df_test.shape)
```

Output:
```
Train DataFrame: (891, 11)
Test DataFrame:  (418, 10)
```

## Check missings

```python
# df.isnull().sum()
```

```python
# df_test.isnull().sum()
```

# Exercise 1 (2pts):
Extract the title (Mr, Mrs, ... ) from the "Name" column.

Tips:
- split(',')[1] to get the 2nd part, and remove the surnamename
- split('.')[0] to get the 1str part, and remove the name

<!-- grade_id: cell-d56107842f0f54d2, points: 0, grade: false, solution: true -->
```python
# CODE HERE get_Title_from_Name funtion
# Create this function using lambda (not def)

get_Title_from_Name = lambda name: name.split(',')[1].split('.')[0].strip()
# Your code goes here

df['Title']      = df['Name'].apply(get_Title_from_Name)
df_test['Title'] = df_test['Name'].apply(get_Title_from_Name)
```

<!-- grade_id: cell-f3d989a17018344b, points: 2, grade: true, solution: false -->
```python
assert df['Title'].values[0] == "Mr"
assert df['Title'].values[1] == "Mrs"
assert df['Title'].values[2] == "Miss"

assert df_test['Title'].values[0] == "Mr"
assert df_test['Title'].values[1] == "Mrs"
assert df_test['Title'].values[414] == "Dona"
```

# Exercise 2 (1pts):
Apply the title_dictionary to get a better information about the title. You have to overwrite the Title variable.

```python
title_dictionary = {
    "Capt": "Officer",
    "Col": "Officer",
    "Major": "Officer",
    "Jonkheer": "Royalty",
    "Don": "Royalty",
    "Sir" : "Royalty",
    "Dr": "Officer",
    "Rev": "Officer",
    "the Countess":"Royalty",
    "Mme": "Mrs",
    "Mlle": "Miss",
    "Ms": "Mrs",
    "Mr" : "Mr",
    "Mrs" : "Mrs",
    "Miss" : "Miss",
    "Master" : "Master",
    "Lady" : "Royalty"
}
```

<!-- grade_id: cell-5c6a842c812fafda, points: 0, grade: false, solution: true -->
```python
# Use map to apply the prevous dict
#

df['Title']      = df.Title.map(title_dictionary)
df_test['Title'] = df_test.Title.map(title_dictionary)

# Your code goes here
```

<!-- grade_id: cell-669f5a637e57e835, points: 1, grade: true, solution: false -->
```python
assert df['Title'].values[886] == "Officer"
assert df_test['Title'].values[417] == "Master"
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-83-319aac91d420> in <module>
----> 1 assert df['Title'].values[886] == "Officer"
      2 assert df_test['Title'].values[417] == "Master"

AssertionError:
```

# Exercise OPTINAL (0pts):
Try to extract some information from the feature **Ticket**. Search on Internet if that colum has some kind of information.

```python
df_test['Ticket']
```

Output:
```
PassengerId
892                 330911
893                 363272
894                 240276
895                 315154
896                3101298
               ...        
1305             A.5. 3236
1306              PC 17758
1307    SOTON/O.Q. 3101262
1308                359309
1309                  2668
Name: Ticket, Length: 418, dtype: object
```

# Exercise OPTIONAL (0pts):
Try to extract some information from the feature **Cabin**. Search on Internet if that colum has some kind of information.

```python
print(df_test['Cabin'].isna().sum())
```

Output:
```
327
```

# Preprocessing
For X data, notice that...
- We drop Survived because is the target variable
- We drop Name because we have extracted the Title: Mr, Mrs, ...
- We drop Ticket because it has no information -> see df.Ticket.nunique()
- We drop Cabin because it has a lot of missings (77% are missings)

Then, we identify **numerical** variables and **categorical** variables,

```python
x = df.drop(columns=["Survived", 'Name', 'Ticket', 'Cabin']) # X DATA (WILL BE TRAIN+VALID DATA)
y = df["Survived"] # 0 = No, 1 = Yes

x_test = df_test.drop(columns=['Name', 'Ticket', 'Cabin']) # # X_TEST DATA (NEW DATA)
```

```python
cat_vars  = ['Sex', 'Embarked', 'Title']         # x.select_dtypes(include=[object]).columns.values.tolist()
num_vars  = ['Pclass', 'SibSp', 'Parch', 'Fare', 'Age'] # x.select_dtypes(exclude=[object]).columns.values.tolist()

print("\nNumerical features:\n", num_vars)
print("\nCategorical features:\n", cat_vars)
```

Output:
```
Numerical features:
 ['Pclass', 'SibSp', 'Parch', 'Fare', 'Age']

Categorical features:
 ['Sex', 'Embarked', 'Title']
```

# Exercise 3 (2pts):
Create a **ColumnTransformer for Tree Models**. You need to create 2 pipelines (one for numerical and other for categories). Remember:
- Categorical pipeline: Some SimpleImputer -> Some Encoder
- Numerical pipeline: Some SimpleImputer -> NO Encoder

<!-- grade_id: cell-c607e75fb38ea248, points: 0, grade: false, solution: true -->
```python
'''
num_preprocessing = pipeline.Pipeline(steps=[ ])

cat_preporcessing = pipeline.Pipeline(steps=[
  # Some SimpleImputer here
  # Some Encoder here. Remember to handle_unknown
])
'''
### BEGIN SOLUTION
num_4_treeModels = pipeline.Pipeline(steps=[
    ('imputer', impute.SimpleImputer(strategy='mean')),
])

cat_4_treeModels = pipeline.Pipeline(steps=[
    ('imputer', impute.SimpleImputer(strategy='constant', fill_value='missing')),
    ('ordinal', preprocessing.OrdinalEncoder(handle_unknown='use_encoded_value', unknown_value = 10)) #  ONLY IN VERSION 0.24
])

tree_prepro = compose.ColumnTransformer(transformers=[
    ('num', num_4_treeModels, num_vars),
    ('cat', cat_4_treeModels, cat_vars),
], remainder='drop') # Drop other vars not specified in num_vars or cat_vars


tree_prepro
### END SOLUTION

# Your code goes here
```

Output:
```
ColumnTransformer(transformers=[('num',
                                 Pipeline(steps=[('imputer', SimpleImputer())]),
                                 ['Pclass', 'SibSp', 'Parch', 'Fare', 'Age']),
                                ('cat',
                                 Pipeline(steps=[('imputer',
                                                  SimpleImputer(fill_value='missing',
                                                                strategy='constant')),
                                                 ('ordinal',
                                                  OrdinalEncoder(handle_unknown='use_encoded_value',
                                                                 unknown_value=10))]),
                                 ['Sex', 'Embarked', 'Title'])])
```

<!-- grade_id: cell-e636558135fb5975, points: 2, grade: true, solution: false -->
```python
assert type(tree_prepro)      is compose._column_transformer.ColumnTransformer
assert type(num_4_treeModels) is pipeline.Pipeline
assert type(cat_4_treeModels) is pipeline.Pipeline
assert len(num_4_treeModels) == 1
assert len(cat_4_treeModels) == 2
```

# Exercise 4 (1pts):
1. Complete the diccionary with some Tree Models.
2. Then we put each model in a Pipeline where:
   - first is the prepocessing with the column Transformer
   - Then is the Tree model
3. Display the fullpipeline of the LGBMClassifier

```python
from sklearn.tree          import DecisionTreeClassifier
from sklearn.ensemble      import ExtraTreesClassifier
from sklearn.ensemble      import RandomForestClassifier
from sklearn.ensemble      import AdaBoostClassifier
from sklearn.ensemble      import GradientBoostingClassifier
from sklearn.experimental  import enable_hist_gradient_boosting # Necesary for HistGradientBoostingClassifier
from sklearn.ensemble      import HistGradientBoostingClassifier
from xgboost               import XGBClassifier
from lightgbm              import LGBMClassifier
from catboost              import CatBoostClassifier
```

<!-- grade_id: cell-76a3be8223730c5a, points: 0, grade: false, solution: true -->
```python

tree_classifiers = {
  "Decision Tree": DecisionTreeClassifier(),
  "Extra Trees": ExtraTreesClassifier(n_estimators=100),
  "Random Forest": RandomForestClassifier(n_estimators=100),
  "AdaBoost": AdaBoostClassifier(n_estimators=100),
  "Skl GBM": GradientBoostingClassifier(n_estimators=100),
  "Skl HistGBM": HistGradientBoostingClassifier(max_iter=100),
  "XGBoost": XGBClassifier(n_estimators=100),
  "LightGBM": LGBMClassifier(n_estimators=100),
  "CatBoost": CatBoostClassifier(n_estimators=100)
}


tree_classifiers = {name: pipeline.make_pipeline(tree_prepro, model) for name, model in tree_classifiers.items()}

tree_classifiers["LightGBM"]
```

Output:
```
Pipeline(steps=[('columntransformer',
                 ColumnTransformer(transformers=[('num',
                                                  Pipeline(steps=[('imputer',
                                                                   SimpleImputer())]),
                                                  ['Pclass', 'SibSp', 'Parch',
                                                   'Fare', 'Age']),
                                                 ('cat',
                                                  Pipeline(steps=[('imputer',
                                                                   SimpleImputer(fill_value='missing',
                                                                                 strategy='constant')),
                                                                  ('ordinal',
                                                                   OrdinalEncoder(handle_unknown='use_encoded_value',
                                                       
```

<!-- grade_id: cell-d4744022b37e2e9b, points: 1, grade: true, solution: false -->
```python
for pipe in tree_classifiers.values():
    assert type(pipe) is pipeline.Pipeline
```

# Exercise 5 (3pts):
Define a simple split validation strategy with:
- 80% for train
- 20% for validation
- With stratification
- random_state=0

And train all the models in a for loop

<!-- grade_id: cell-64f17e0d448bca7e, points: 0, grade: false, solution: true -->
```python

x_train, x_val, y_train, y_val = model_selection.train_test_split(x, y, test_size= 0.2, stratify = y, random_state = 0 )


results = pd.DataFrame({'Model': [], 'Accuracy': [], 'Bal Acc.': [], 'Time': []})


for model_name, model in tree_classifiers.items():
    start_time = time.time()
    
    #print(type(model))
    # FOR EVERY PIPELINE (PREPRO + MODEL) -> TRAIN WITH TRAIN DATA (x_train)
    model.fit(x_train, y_train)
    
    # GET PREDICTIONS USING x_val
    pred = model.predict(x_val)

    total_time = time.time() - start_time

    results = results.append({"Model":    model_name,
                              "Accuracy": metrics.accuracy_score(y_val, pred)*100,
                              "Bal Acc.": metrics.balanced_accuracy_score(y_val, pred)*100,
                              "Time":     total_time},
                              ignore_index=True)
                                                        

# Your code goes here

results_ord = results.sort_values(by=['Accuracy'], ascending=False, ignore_index=True)
results_ord.index += 1 
#results_ord.style.bar(subset=['Accuracy', 'Bal Acc.'], vmin=0, vmax=100, color='#5fba7d')
results_ord
```

Output:
```
C:\Users\Jon\.conda\envs\NLP\lib\site-packages\xgboost\sklearn.py:888: UserWarning: The use of label encoder in XGBClassifier is deprecated and will be removed in a future release. To remove this warning, do the following: 1) Pass option use_label_encoder=False when constructing XGBClassifier object; and 2) Encode your labels (y) as integers starting with 0, i.e. 0, 1, 2, ..., [num_class - 1].
  warnings.warn(label_encoder_deprecation_msg, UserWarning)
[16:11:33] WARNING: ..\src\learner.cc:1061: Starting in XGBoost 1.3.0, the default evaluation metric used with the objective 'binary:logistic' was changed from 'error' to 'logloss'. Explicitly set eval_metric if you'd like to restore the old behavior.
Learning rate set to 0.073611
0:	learn: 0.6698078	total: 29.8ms	remaining: 2.95s
1:	learn: 0.6443828	total: 31.4ms	remaining: 1.54s
2:	learn: 0.6218994	total: 32.4ms	remaining: 1.05s
3:	learn: 0.6043828	total: 32.9ms	remaining: 789ms
4:	learn: 0.5864178	total: 33.7ms	remaining: 641ms
5:	lea
```

<!-- grade_id: cell-af27b33fe88ad384, points: 3, grade: true, solution: false -->
```python
assert results_ord["Accuracy"].min() > 75
assert results_ord["Bal Acc."].min() > 70
assert len(results_ord) == 9
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-112-b57998e2e157> in <module>
----> 1 assert results_ord["Accuracy"].min() > 75
      2 assert results_ord["Bal Acc."].min() > 70
      3 assert len(results_ord) == 9

AssertionError:
```

# Exercise 6 (3pts):
Define a 10 Fold cross validation strategy with:
- With stratification
- shuffle=True
- random_state=0

And train all the models in a for loop.

Tip you can use **[cross_val_predict](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.cross_val_predict.html)** for both training and predict with 

<!-- grade_id: cell-37bf16dc25b43732, points: 0, grade: false, solution: true -->
```python
"""
skf = model_selection.StratifiedKFold(
    # CODE HERE
)
"""
# Your code goes here

results = pd.DataFrame({'Model': [], 'Accuracy': [], 'Bal Acc.': [], 'Time': []})

"""
for model_name, model in tree_classifiers.items():
    start_time = time.time()
        
    # TRAIN AND GET PREDICTIONS USING cross_val_predict() and x,y
    pred = # CODE HERE

    total_time = time.time() - start_time

    results = results.append({"Model":    model_name,
                              "Accuracy": metrics.accuracy_score(y_val, pred)*100,
                              "Bal Acc.": metrics.balanced_accuracy_score(y_val, pred)*100,
                              "Time":     total_time},
                              ignore_index=True)
                              
                              
"""

# Your code goes here


results_ord = results.sort_values(by=['Accuracy'], ascending=False, ignore_index=True)
results_ord.index += 1 
results_ord.style.bar(subset=['Accuracy', 'Bal Acc.'], vmin=0, vmax=100, color='#5fba7d')
```

Output:
```
<pandas.io.formats.style.Styler at 0x297aa7b8d30>
```

<!-- grade_id: cell-57cc085faad513cf, points: 3, grade: true, solution: false -->
```python
assert results_ord["Accuracy"].min() > 75
assert results_ord["Bal Acc."].min() > 75
assert len(results_ord) == 9
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-107-50685871c2e2> in <module>
----> 1 assert results_ord["Accuracy"].min() > 75
      2 assert results_ord["Bal Acc."].min() > 75
      3 assert len(results_ord) == 9

AssertionError:
```

# Exercise 7.1
Train with all data the best model

```python
# best_model = tree_classifiers["SELECT MY BEST MODEL HERE"]

# Fit best model with all data


# Your code goes here
```

# Exercise 7.2 (2pts)
With your best model, generate the predicitions for test data (x_test)

<!-- grade_id: cell-44b93a0dbd4eb6fc, points: 0, grade: false, solution: true -->
```python
# test_pred = # Get the predictions for x_test

# Your code goes here
```

<!-- grade_id: cell-a29fc8d24284520e, points: 2, grade: true, solution: false -->
```python
assert len(test_pred) == 418
assert np.unique(test_pred).tolist() == [0,1]
```

# Exercise 7.3

Submit to kaggle.

- You can use the kaggle command line app. Check https://github.com/Kaggle/kaggle-api

```python
sub = pd.DataFrame(test_pred, index=x_test.index, columns=["Survived"])
sub.head()
```

```python
sub.to_csv("sub.csv")
```

```python
!kaggle competitions submit -c titanic -f sub.csv -m "My submission message"
```